# Entrega 2
La entrega pasada nos enfocamos en realizar una primera aproximación. Ahora, a no ser que se trabaje con Deep Learning, para hacer un trabajo de clasificación, es necesario tener 2 elementos principales: 
>un descriptor \
>un clasificador

El primero, tal como su nombre lo indica, es el que nos va a permitir describir la imagen, es decir expresar sus características de forma cuantitativa para poder tomar decisiones sobre esta, ese es el vector de TM que hicimos en la entrega pasada. El segundo, cuyo nombre también nos da una pista de su utilidad, sirve para tomar las características que nos da el descriptor y poder asignar clases a las imágenes basándonos en estas. En esta entrega nos enfocaremos en la realización de los descriptores de patrones visuales que hemos visto en la magistral usando como clasificador el vecino mas cercano que hicimos el laboratorio pasado.

## Parte 1: Modificaciones al vecino mas cercano
Lo primero que vamos a hacer será modificar nuestra clase del vecino mas cercano que hicimos la entrega pasada. Esto es debido a que en la entrega pasada en el interior de la funcion tanto de entrenamiento como de predicción se calculaban los descriptores. Sin embargo haremos unas modificaciones a este metodo con el fin de que no se tengan que calcular los descriptores cada vez que se quiera clasificar. 
> **Nota**: los cambios que tienen que hacer a sus codigos son minimos, unicamente es eliminar la linea que calculaba el descriptor y asumir que X es directamente el descriptor


<!-- grade_id: cell-25b8e7de9fa3382a, points: 0, grade: false, solution: true -->
```python
class NN_classifier:
    def __init__(self,descriptor):
        '''
        :param descriptor (dict): Diccionario de templates de entrenamiento.
        '''
        self.vectores='' #atributo del vector de descriptores de entrenamiento
        self.etiquetas='' #atributo del vector de etiquetas
        # YOUR CODE HERE
        raise NotImplementedError()
    def fit(self,X,y):
        '''
        :param X: Imagenes de entrenamiento
        :param y: Labels de las imagenes
        :return: self (Trained descriptor)
        '''
        # YOUR CODE HERE
        raise NotImplementedError()
        return self
    
    def predict(self,X):
        '''
        :param X (array): Imagenes a predecir
        :return: Arreglo de etiquetas predichas
        '''
        # YOUR CODE HERE
        raise NotImplementedError()
        return etiqueta
```

## Parte 2: Textura
Una de las formas en las que se puede hacer una descripción de la textura de una imagen es mediante el histograma de textones, el cual es una forma de cuantificar la presencia de las textura existentes. Para esto, debemos implementar una función que calcule dicho histograma utilizando un diccionario de textones. Por lo anterior, implementaremos tres funciones auxiliares que nos permitirán obtener sus descriptores de textura. Estas funciones deben realizarlas para poder calcular la respuesta a los filtros, calcular el diccionario de textones y calcular el histograma de textones.

Primeramente vamos a sacar los filtros de interes, los cuales están en el archivo filterbank.mat


<!-- grade_id: cell-ac1e025aa0c5be43, points: 0, grade: false, solution: true -->
```python
filterbank = '' #Variable que contendrá los filtros para realizar los textones
image = '' #Variable que contendrá una imagen a color para probar las funciones
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ef968a6d21127fdd, points: 0.1, grade: true, solution: false -->
```python
import numpy as np
assert not type(filterbank) == dict, f'su banco debería ser una matriz, no un diccionario, revise los elementos y busque el de interes'
print(f'la forma del banco de filtros es: {filterbank.shape}')
assert filterbank.shape == (49,49,18), f'su banco de filtros debería ser de tamaño (49,49,18), no {filterbank.shape}'
```

### 2.1 calcular la respuesta a los filtros
Esta primera función tiene como propósito hacer el cálculo de las respuestas a los filtros. Como podrán recordar de lo visto en la magistral, el primer paso para obtener sus textones, así como el mapa de textones de una imagen, es calcular la respuesta de la imagen en escala de grises a cada uno de los filtros que se encuentren en el banco. Finalmente, deben almacenar la respuesta a los filtros de cada uno de los píxeles de la imagen. 
> **Nota**: recuerde que el procedimiento debe ser hecho en escala de grises. Además, vamos a querer que cada pixel sea un dato, por lo que no usaremos una matriz de descriptores sino una lista de ellos.

<!-- grade_id: cell-aabff9e04b2e4b36, points: 0, grade: false, solution: true -->
```python
def calculate_filter_response(img,filters):
    '''
    :param img: imagen a color de la cual se va a calcular la respuesta a los filtros
    :param filters: filtros que serán usados
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
    return np.array(flattened_response)
   
```

<!-- grade_id: cell-4252fd21b88111d3, points: 0.3, grade: true, solution: false -->
```python
pixel_response = calculate_filter_response(image,filterbank)
assert len(pixel_response.shape) == 2, f'la respuesta al filtro deberia tener 2 dimensiones, no {len(pixel_response.shape)}'
assert pixel_response.shape == (image.shape[0]*image.shape[1],18),f'la respuesta debería ser de tamaño {(image.shape[0]*image.shape[1],18)} no {pixel_response.shape}'
```

### 2.2 Calcular el diccionario de textones
Esta segunda función tiene como propósito construir un diccionario de textones. Esta deberá recibir por parámetro el arreglo de píxeles descritos por sus respuestas a los filtros (en las imágenes de entrenamiento). Luego, haciendo uso de un algoritmo de clusterización (comúnmente k-means), agruparlos en lo que serán "tipos de textura". Por último, los centroides de los clusters de textura generados, los cuales serán nuestros textones, deberán ser retornados.

<!-- grade_id: cell-52167af5ed2da1ff, points: 0, grade: false, solution: true -->
```python
def calculate_texton_dictionary(pixels,n_textons):
    '''
    :param img:
    :param filters:
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
    return centroids # este debe el modelo KMeans
```

<!-- grade_id: cell-28aee991d64bd81d, points: 0.3, grade: true, solution: false -->
```python
texton_dictionary = calculate_texton_dictionary(pixel_response,10)
try:
    texton_dictionary.n_features_in_
except:
    assert False,f'Su modelo no fue entrenado'
assert texton_dictionary.n_features_in_ == 18,f'su modelo tiene un espacio de {texton_dictionary.n_features_in_} dimensiones, no 18 como debería'
assert len(texton_dictionary.labels_)==image.shape[0]*image.shape[1],f'su modelo tiene {len(texton_dictionary.labels_)} puntos, no {image.shape[0]*image.shape[1]}'
assert len(texton_dictionary.cluster_centers_)==10,f'su modelo KMeans fue entrenado con {len(texton_dictionary.cluster_centers_)} clusters, no 10'
```

### 2.3 Calular histograma de textones
Esta última función generará el mapa de textones de una imagen, en base al cual obtendrán como respuesta el histograma de textones asociado. Esta deberá recibir por parámetro el arreglo de píxeles de la imagen a describir (la respuestas a los filtros) y asociar cada píxel a la etiqueta del centroide (textón) del cluster correspondiente, para finalmente construir y retornar el histograma de textones.
> **nota**: recuerde normalizar el histograma. Y recuerde que aqui tiene un numero finito de valores completos, por lo que debe tener un bin por cada valor

<!-- grade_id: cell-6b0dea1c029f476a, points: 0, grade: false, solution: true -->
```python
def calculate_texton_histogram(pixels,textons):
    '''
    :param pixels: Array de respuesta a los filtros de la imagen 
    :param texton_dictionary: diccionario de textones(Modelo KMeans)
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
    return texton_histogram 
```

<!-- grade_id: cell-17700f298b97879b, points: 0.3, grade: true, solution: false -->
```python
demo_hist = calculate_texton_histogram(pixel_response,texton_dictionary)
assert len(demo_hist)==10, f'el histograma debería tener 10 bins, uno por centroide, no {len(demo_hist)}'
assert np.isclose(np.sum(demo_hist),1,0.05), f'la suma del histogtama debería ser 1, no {np.sum(demo_hist)}'
```

Este es el unico descriptor que tendremos que entrenar, ya que los demás no dependen sino de las caracteristicas propias de la imagen. Por ello deben recalcular el diccionario de textones usando almenos 1 imagen de cada clase. 

```python
texton_dictionary = '' #Variable que tendrá el diccionario de textones a usar. Usen uno de 10 textones como en el ejemplo
```

### 2.4 Funcion de calculo de descriptores 
Como ya fue mencionado, haciendo uso de las funciones previamente mencionadas y el/los diccionarios de textones que han generado y almacenado, deberán construir una función que genere el descriptor de textura. Un punto importante que deben tener en cuenta es que usarán esta función no solo para calcular los descriptores, sino para almacenarlos. La razón de esto es que los descriptores en algunos casos pueden tardar un tiempo considerable en ser calculados. Debido a lo anterior, sería computacionalmente costoso y tomaría tiempo en exceso hacer el cálculo de los descriptores en cada una de las entregas. El archivo que deben almacenar debe contener dos listas las cuáles tengan posiciones correspondientes entre ellas, una con la información de los descriptores (uno por imagen) y la otra con la anotación correspondiente. La función debe tener la siguiente estructura:

<!-- grade_id: cell-ee86f35b0ee122b8, points: 0, grade: false, solution: true -->
```python
def texture(images, labels, route, textons):
    '''
    :param images: Imágenes de la cuales se calculará el descriptor.
    :param labels: Etiquetas de las imágenes.
    :param route: Ruta donde serán guardados los descriptores junto con las etiquetas.
    :param textons: Diccionario de textones o ruta al archivo que lo contiene.
    :return features: Arreglo de etiquetas predichas con los descriptores de las imágenes (el mismo que se almacenará en route)
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
    return np.array(features)
```

<!-- grade_id: cell-5f0e59b1092e44dd, points: 0, grade: false, solution: true -->
```python
train_images = '' #vector de imagenes de entrenamiento
train_annots = '' #vector de anotaciones de entrenamiento
train_texture_route = '' #ruta donde quieren guardar sus descriptores y labels

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6e4171c1bd588d65, points: 0.5, grade: true, solution: false -->
```python
import os
texture_desc = texture(train_images,train_annots,train_texture_route,texton_dictionary)
assert len(texture_desc.shape) == 2 ,f'Su vector de descriptores debe tener 2 dimensiones, no len(texture_desc.shape)'
assert texture_desc.shape[0] == len(train_images),f'no existe un decriptor por imagen, hay {texture_desc.shape[0]} descriptores y {len(train_images)} imagenes'
assert not np.isclose(np.std(texture_desc),0,0.001),f'su funcion retorna descriptores similares o identicos'
assert os.path.isfile(train_texture_route),'no se está guardando el archivo'

assert texture_desc.shape[1] ==10,f'su modelo KMeans fue entrenado con {len(texton_dictionary.cluster_centers_)} clusters, no 10'
```

## Parte 3: Color
Lo primero que se debe realizar, como ya se dijo, es el descriptor mediante el cual vamos a representar una imagen en un vector. En esta entrega se trabajará únicamente con descriptores basados en color. Estos, se basan en representar la imagen según los diferentes colores que la compongan mediante el uso de histogramas, ya que pretende caracterizar la frecuencia de cada una de las intensidades de los colores.

Este descriptor de histogramas de color será el primero que usaremos. La función debe tener la capacidad de calcular tanto histogramas de color conjuntos como histogramas de color concatenados los cuáles pueden obtenerlos utilizando las funciones ```CatColorHistogram``` y ```JointColorHistogram``` que se encuentran a continuación. Además, debe poder utilizar diferentes canales de color como ```rgb```, ```hsv``` y ```lab```. 

```python
# Functions for color histograms

def CatColorHistogram(img, num_bins, min_val=None, max_val=None):
    """
    Calculate concatenated histogram for color images
    By: Natalia Valderrama built on Maria Fernanda Roa's code
    
    Arguments: img (numpy.array) -- 2D color image
    num_bins (array like of ints) -- Number of bins per channel.
    If an int is given, all channels will have same amount of bins.

    Keyword Arguments:
    min_val (array like of ints) -- Minimum intensity range value per channel
    If an int is given, all channels will have same minimum. (default: {None})
    max_val (array like of ints) -- Maximum intensity range value per channel
    If an int is given, all channels will have same maximum. (default: {None})

    Returns: [numpy.array] -- Array containing concatenated color histogram of size num_bins*3.
    """
    assert len(img.shape) == 3, 'img must be a color 2D image'

    #Transform image to float dtype
    img = img_as_float(img)
    _, _, n_channels = img.shape

    #Verify input parameters
    assert isinstance(num_bins, (int, tuple, list, np.array)),'num_bins must be int or array like'

    if isinstance(num_bins, int):
        num_bins = np.array([num_bins]*n_channels)
    else:
        num_bins = np.array(num_bins)

    assert len(num_bins) == n_channels,'num_bins length and number of channels differ'
 
    if min_val is None:
        min_val = np.min(img, (0,1))
    else:
        assert isinstance(min_val, (int, tuple, list, np.array)),'min_val must be int or array like'
        if isinstance(min_val, int):
            min_val = np.array([min_val]*n_channels)
        else:
            min_val = np.array(min_val)

    assert len(min_val) == n_channels,'min_val length and number of channels differ'
    
    min_val = min_val.reshape((1, 1, -1))

    if max_val is None:
        max_val = np.max(img, (0,1))
    else:
        assert isinstance(max_val, (int, tuple, list, np.array)),'max_val must be int or array like'
        if isinstance(max_val, int):
            max_val = np.array([max_val]*n_channels)
        else:
            max_val = np.array(max_val)

    assert len(max_val) == n_channels,'max_val length and number of channels differ'
    max_val = max_val.reshape((1, 1, -1)) + 1e-5
    concat_hist = np.zeros(np.sum(num_bins), dtype=np.int)
    # Scale intensities (intensities are scaled within the range for each channel)
    # Values now are between 0 and 1
    img = (img - min_val) / (max_val - min_val)
    sum_value = 0

    for c in range(n_channels):
        # Calculate index matrix for each channel

        idx_matrix = np.floor(img[...,c]*num_bins[c]).astype('int')
        idx_matrix = idx_matrix.flatten() + sum_value
        sum_value += num_bins[c]
        
        #Create concatenated histogram
        for p in range(len(idx_matrix)):
            concat_hist[idx_matrix[p]] += 1
    
    return concat_hist/np.sum(concat_hist)

def JointColorHistogram(img, num_bins, min_val=None, max_val=None):
    """
    Calculate joint histogram for color images
    By: Maria Fernanda Roa

    Arguments: img (numpy.array) -- 2D color image
    num_bins (array like of ints) -- Number of bins per channel.
    If an int is given, all channels will have same amount of bins.

    Keyword Arguments:
    min_val (array like of ints) -- Minimum intensity range value per channel
    If an int is given, all channels will have same minimum. (default: {None})
    max_val (array like of ints) -- Maximum intensity range value per channel
    If an int is given, all channels will have same maximum. (default: {None})

    Returns: [numpy.array] -- Array containing joint color histogram of size num_bins.
    """

    assert len(img.shape) == 3, 'img must be a color 2D image'

    #Transform image to float dtype
    img = img_as_float(img)
    _, _, n_channels = img.shape

    #Verify input parameters
    assert isinstance(num_bins, (int, tuple, list, np.array)),'num_bins must be int or array like'

    if isinstance(num_bins, int):
        num_bins = np.array([num_bins]*n_channels)
    else:
        num_bins = np.array(num_bins)

    assert len(num_bins) == n_channels,'num_bins length and number of channels differ'
 
    if min_val is None:
        min_val = np.min(img, (0,1))
    else:
        assert isinstance(min_val, (int, tuple, list, np.array)),'min_val must be int or array like'
        if isinstance(min_val, int):
            min_val = np.array([min_val]*n_channels)
        else:
            min_val = np.array(min_val)

    assert len(min_val) == n_channels,'min_val length and number of channels differ'
    
    min_val = min_val.reshape((1, 1, -1))

    if max_val is None:
        max_val = np.max(img, (0,1))
    else:
        assert isinstance(max_val, (int, tuple, list, np.array)),'max_val must be int or array like'
        if isinstance(max_val, int):
            max_val = np.array([max_val]*n_channels)
        else:
            max_val = np.array(max_val)

    assert len(max_val) == n_channels,'max_val length and number of channels differ'
    max_val = max_val.reshape((1, 1, -1)) + 1e-5

    joint_hist = np.zeros(num_bins, dtype=np.int)
    num_bins = num_bins.reshape((1, 1, -1))

    # Scale intensities (intensities are scaled within the range for each channel)
    # Values now are between 0 and 1
    img = (img - min_val) / (max_val - min_val)

    # Calculate index matrix
    idx_matrix = np.floor(img*num_bins).astype('int')
    idx_matrix = idx_matrix.reshape((-1, n_channels))
    
    #Create joint histogram
    for p in range(len(idx_matrix)):
        joint_hist[tuple(idx_matrix[p, :])] += 1
        
    return joint_hist/np.sum(joint_hist)
```

<!-- grade_id: cell-5378cc00e581e1ab, points: 0, grade: false, solution: true -->
```python
def color(images, labels, route, type, space_bins, color_space):
    '''
    :param images}: Lista con las imágenes (o rutas de las imágenes) a la cuales se les calculará el descriptor.
    :param labels: Lista con las etiquetas de las imágenes pasadas por parámetro.
    :param route: Ruta donde serán guardados los descriptores junto con las etiquetas.
    :param type: Parámetro que indica el tipo de histograma de color que se calculará, puede ser ``joint'' o ``concat''.
    :param space_bins: Número entero que indica el numero de bins que existirá por cada canal de color.
    :param color_space: Parámetro que indica el espacio de color que se usará, puede ser ``rgb'', ``hsv'', o ``lab''.
    
    :return features: Arreglo de etiquetas predichas con los descriptores de las imágenes (el mismo que se almacenará en route)
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
    return np.array(features)
```

<!-- grade_id: cell-ac2618f5c7510774, points: 0, grade: false, solution: true -->
```python
train_color_route = ''
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9638e0c60035ebd3, points: 0.5, grade: true, solution: false -->
```python
import os
color_desc = color(train_images,train_annots,train_color_route,"joint",5,"rgb")
assert len(color_desc.shape) == 2 ,f'Su vector de descriptores debe tener 2 dimensiones, no {len(color_desc.shape)}'
assert color_desc.shape[0] == len(train_images),f'no existe un decriptor por imagen, hay {color_desc.shape[0]} descriptores y {len(train_images)} imagenes'
assert not np.isclose(np.std(color_desc),0,0.001),f'su funcion retorna descriptores similares o identicos'
assert os.path.isfile(train_color_route),'no se está guardando el archivo'

assert color_desc.shape[1] ==5**3,f'su modelo histograma conjunto tiene {color_desc.shape[1]} bins, no {5**3}'
color_desc_cat = color(train_images,train_annots,train_color_route,"concat",5,"rgb")
assert color_desc_cat.shape[1] ==5*3,f'su modelo histograma conjunto tiene {color_desc_cat.shape[1]} bins, no {5*3}'
```

## Parte 4: Forma
Finalmente, el último descriptor que utilizaremos en este miniproyecto será basado en forma, para esto usarán el {Histograma de gradientes orientados (HOG) que pueden encontrar en el siguiente link.

https://scikit-image.org/docs/dev/api/skimage.feature.html#skimage.feature.hog. 

Este descriptor es especialmente importante ya que es el único descriptor de los que utilizaremos que tiene en cuenta la posición espacial. Para esto, deberán implementar una función que retorne el histograma de gradientes orientados. En esta, deben poder variar por parámetro 2 parámetros que ustedes establezcan para experimentar. Al igual que las funciones de los otros descriptores, esta función no solo debe retornar los descriptores de las imágenes sino que debe poder almacenarlos en un archivo local.  Esta función debe tener la siguiente estructura:\\

<!-- grade_id: cell-407677d678c13cae, points: 0, grade: false, solution: true -->
```python
def shape(images, labels, route, param1 = 0, param2 = 0):
    '''
    :param images: Imágenes de la cuales se calculará el descriptor.
    :param labels: Etiquetas de las imágenes.
    :param route: Ruta donde serán guardados los descriptores junto con las etiquetas.
    :param param1: Parámetro 1 que variarán para calcular HOG. Debe ser por defecto un valor por las pruebas
    :param param2: Parámetro 2 que variarán para calcular HOG.Debe ser por defecto un valor por las pruebas
    
    :return features: Arreglo de etiquetas predichas con los descriptores de las imágenes (el mismo que se almacenará en route)
    '''

    # YOUR CODE HERE
    raise NotImplementedError()
    return np.array(features)
```

<!-- grade_id: cell-70f77ee5f2d76d29, points: 0, grade: false, solution: true -->
```python
train_shape_route = ''
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2645d62759258264, points: 0.5, grade: true, solution: false -->
```python
import os
shape_desc = shape(train_images,train_annots,train_shape_route)
assert len(shape_desc.shape) == 2 ,f'Su vector de descriptores debe tener 2 dimensiones, no {len(shape_desc.shape)}'
assert shape_desc.shape[0] == len(train_images),f'no existe un decriptor por imagen, hay {shape_desc.shape[0]} descriptores y {len(train_images)} imagenes'
assert not np.isclose(np.std(shape_desc),0,0.001),f'su funcion retorna descriptores similares o identicos'
assert os.path.isfile(train_shape_route),'no se está guardando el archivo'
```

## Parte 5: Experimentación
Ahora que ya tenemos implementados nuestros nuevos descriptores de patrones visuales vamos a experimentar con ellos, usando

### 5.1 Textura:
Deben realizar la experimentación variando el número de textones que hay en el diccionario (es decir el numero de bins del histograma). Deben construir por lo menos 4 diccionarios diferentes.**(4 experimentos)**

> Nota: La respuesta a los filtros solo deberían calcularla una vez, guarden esto en archivos y calculen los 4 diccionarios e histogramas usando estos.

### 5.2 Color:
Deben realizar la experimentación en grid 2x3, variando el tipo de histograma que se calcula y el espacio de color respectivamente. **(6 experimentos)**

### 5.3 Forma: 
Deben realizar la experimentación en grid 2x2(es decir las 4 combinaciones de 2 parametros con 2 valores cada uno) escogiendo 2 parámetros de la función para calcular HOG y 2 valores para cada uno de dichos parámetros. **(4 experimentos)**



```python
# use este espacio para hacer sus experimentos
```

## Articulo

Una vez realizado el procedimiento es fundamental que se continue con la construcción del articulo de forma progresiva, por lo cual incluiremos lo que hicimos en esta entrega.

### Abordaje

#### Descriptores
Inicialmente explicarán los descriptores, estos no solo los usarán para su método de baseline, sino para el método final. Para esta entrega, describan los 3 descriptores, además de los parámetros que variaron y su propósito. Para esta descripción pueden basarse en las siguientes preguntas:
- ¿En qué consistía su descriptor?
- ¿Qué pretendía mostrar?
- ¿Qué características de la imagen usa?
- ¿Por qué este descriptor sería de utilidad?

#### Baseline
Paso siguiente, deben describir su método de vecino mas cercano, el cual será su *baseline* con el que deberán compararse al tener su metodología experimental.
Para esto describan brevemente el procedimiento que usaron para poder asignar etiquetas basados en clusterización. No incluyan información sobre la librería o la función específica, basta con el método, parámetros y describir su razonamiento. Pueden guiarse de las siguientes preguntas:

- ¿En qué consiste su método?
- ¿De qué elementos y procedimientos se compone?
- ¿Qué ventajas tiene?
- ¿Por qué este clasificador sería de utilidad?

Registren los resultados de sus experimentos del clasificador de *baseline* en su artículo. Incluyan los resultados en una única tabla y resalten en negrilla su mejor resultado, además realice una breve descripción de estos resultados.


### Experimentos
#### Baseline
Finalmente, ya cuentan con los primeros experimentos así como sus primeros resultados, los cuáles deberán ser agregados a su artículo en la sección de experimentos. Primero hagan la descripción de los experimentos que realizaron para los descriptores de color. Luego, incluyan los resultados de la experimentación en una única tabla y señalen en negrilla claramente cuál fue el mejor resultado. Además, deben hacer una breve descripción de sus resultados. Para lo anterior, pueden usar las siguientes preguntas como guía para armar su descripción.
- ¿Cuál fue su mejor método?
- ¿Hubo algún parámetro que influyera mas? (i.e. el histograma conjunto tuvo mejor desempeño en términos generales que el concatenado.)
- ¿Qué combinación no funcionó de forma correcta?
- ¿Existía alguna tendencia?}

### Discusión

#### Descriptores
Incluyan un análisis de lo obtenido en sus experimentos. Para esto analicen las implicaciones que tienen las métricas obtenidas en los resultados así como los resultados cualitativos, teoricen sobre las razones detrás de estos resultados y cómo propondrían mejorar estas falencias.

#### Baseline
Además, deben discutir sobre el método final que obtuvieron para el baseline (NN con su mejor descriptor) y los resultados de la experimentación que los condujeron a esta decisión, analicen las métricas y resultados cualitativos que lograron obtener.
Para la discusión pueden utilizar las siguientes preguntas como guía: 
- ¿Por qué creen que el mejor experimento dio así?
- Si existe alguna tendencia ¿qué podría significar esta?
- ¿Se debe continuar con el descriptor de mejor desempeño?
- ¿Cuál de los descriptores cree usted que obtendrá mejores resultados con los clasificadores supervisados? ¿Por qué?
- ¿Cree usted que su método de \textit{baseline} es óptimo? ¿Que ventajas/dificultades tiene?
- ¿Qué características comunes tienen aquellas imágenes mal clasificadas?