# Miniproyecto 3: Clasificación de animales

## Entrega 3

En la vision artificial, el problema de la clasificación de imagenes tiene 2 partes importantes, el descriptor y el clasificador. En las anteriores entregas ya hemos experimentado con los descriptores, una forma de explicarle a un computador las caracteristicas visuales de una imagen mediante un vector. En esta ocasión, usaremos una unica forma de describir nuestras imagenes y experimentaremos con los diferentes clasificadores que hay, SVM, Random Forest y MLP (Multi Layer Perceptron).

### Parte 1: Carga de descriptores

Lo que debemos hacer primero es volver nuestras imagenes vectores que nuestros computadores puedan entender. En la entrega anterior pasamos todas las imagenes de la base de datos a sus respectivos descriptores, esto es costoso computacionalmente por lo que cargaremos estos vectores para no tener que volverlos a sacar. Utilice únicamente los descritpores que mejores resultados obtuvo de la practica anterior.

<!-- grade_id: cell-8cb66375a49cd51b, points: 0, grade: false, solution: true -->
```python
import numpy as np
import os

num_imagenes='' #numero de imagenes de entrenamiento

# Cargue aqui sus descriptores con sus respectivas labels. Asegurese de que las variables COINSIDAN. Si usted guardo primero las labels y luego los descriptores, cambie el orden
descript_color_train,labels_color_train='','' #Lista de descriptores de color, mejor segun experimentación
descript_shape_train,labels_shape_train='','' #Lista de descriptores de forma, mejor segun experimentación
descript_texture_train,labels_texture_train='','' #Lista de descriptores de textura, mejor segun experimentación

# YOUR CODE HERE

def load_descriptors_labels(route_name):
    route = os.path.join(route_name)
    with open(route, 'rb') as f:
        data = np.load(f, allow_pickle=True).item()
    descriptors = data['descriptors']
    labels = data['labels']
    return descriptors, labels

# el mejor fue el de 17 textones (exp. 4)
descript_texture_train, labels_texture_train = load_descriptors_labels('train_texture_4.npy')
# el mejor fue el de joint+lab (exp. 3)
descript_color_train, labels_color_train = load_descriptors_labels('train_color_3.npy')
# el mejor fue el de 14 orientaciones y 4 pixeles_por_celda (exp. 1)
descript_shape_train, labels_shape_train = load_descriptors_labels('train_shape_1.npy')

num_imagenes = len(descript_shape_train)
```

<!-- grade_id: cell-72a8d9105528c7a7, points: 0.4, grade: true, solution: false -->
```python
assert len(descript_color_train)==len(descript_shape_train), f'Deberia tener la misma cantidad de vectores en todos los descriptores'
assert len(descript_shape_train)==len(descript_texture_train), f'Deberia tener la misma cantidad de vectores en todos los descriptores'
assert len(descript_color_train)==num_imagenes, f'La cantidad de imagenes en la matriz de descriptores no coinside con el numero de imagenes que usted indico'
assert len(descript_color_train)==len(labels_color_train), 'La cantidad de labels deberian ser iguales que la cantidad de vectores de descriptores'
```

### Parte 2: SVM

Ahora, utilizaremos un clasificador de SVM para clasificar nuestras imagenes usando los diferentes descriptores. Inicialice 4 modelos para cada descriptor (12 modelos) variando el valor C y el Kernel. Entrene sus respectivos modelos.

<!-- grade_id: cell-16d711d8932cc29e, points: 0, grade: false, solution: true -->
```python
import numpy as np
from sklearn.svm import SVC

C1='' #valor 1 para C
C2='' #valor 2 para C

kernel1='' #kernel 1
kernel2='' #kernel 2

color_model1='' #kernel 1 y C1
color_model2='' #kernel 1 y C2
color_model3='' #kernel 2 y C1
color_model4='' #kernel 2 y C2

shape_model1='' #kernel 1 y C1
shape_model2='' #kernel 1 y C2
shape_model3='' #kernel 2 y C1
shape_model4='' #kernel 2 y C2

texture_model1='' #kernel 1 y C1
texture_model2='' #kernel 1 y C2
texture_model3='' #kernel 2 y C1
texture_model4='' #kernel 2 y C2

# YOUR CODE HERE
C1 = 0.1
C2 = 10

kernel1 = 'linear'
kernel2 = 'rbf'

color_model1 = SVC(C=C1, kernel=kernel1).fit(descript_color_train, labels_color_train)
color_model2 = SVC(C=C2, kernel=kernel1).fit(descript_color_train, labels_color_train)
color_model3 = SVC(C=C1, kernel=kernel2).fit(descript_color_train, labels_color_train)
color_model4 = SVC(C=C2, kernel=kernel2).fit(descript_color_train, labels_color_train)

shape_model1 = SVC(C=C1, kernel=kernel1).fit(descript_shape_train, labels_shape_train)
shape_model2 = SVC(C=C2, kernel=kernel1).fit(descript_shape_train, labels_shape_train)
shape_model3 = SVC(C=C1, kernel=kernel2).fit(descript_shape_train, labels_shape_train)
shape_model4 = SVC(C=C2, kernel=kernel2).fit(descript_shape_train, labels_shape_train)

texture_model1 = SVC(C=C1, kernel=kernel1).fit(descript_texture_train, labels_texture_train)
texture_model2 = SVC(C=C2, kernel=kernel1).fit(descript_texture_train, labels_texture_train)
texture_model3 = SVC(C=C1, kernel=kernel2).fit(descript_texture_train, labels_texture_train)
texture_model4 = SVC(C=C2, kernel=kernel2).fit(descript_texture_train, labels_texture_train)
```

<!-- grade_id: cell-584d1c49bd193a63, points: 0.7, grade: true, solution: false -->
```python
model1_params_color=color_model1.get_params()
model2_params_color=color_model2.get_params()
model3_params_color=color_model3.get_params()
model4_params_color=color_model4.get_params()

model1_params_shape=shape_model1.get_params()
model2_params_shape=shape_model2.get_params()
model3_params_shape=shape_model3.get_params()
model4_params_shape=shape_model4.get_params()

model1_params_texture=texture_model1.get_params()
model2_params_texture=texture_model2.get_params()
model3_params_texture=texture_model3.get_params()
model4_params_texture=texture_model4.get_params()

assert not texture_model1.fit_status_, 'Su modelo de textura no esta entrenado'
assert not shape_model1.fit_status_, 'Su modelo de forma no esta entrenado'
assert not color_model1.fit_status_, 'Su modelo de color no esta entrenado'
assert model1_params_color['C']==C1, f'El valor de C deberia ser {C1} y es {model1_params_color["C"]}'
assert model2_params_shape['C']==C2, f'El valor de C deberia ser {C2} y es {model2_params_shape["C"]}'
assert model3_params_texture['kernel']==kernel2, f'El kernel deberia ser {kernel2} y es {model3_params_texture["kernel"]}'
assert texture_model4.n_features_in_==len(descript_texture_train[0]), f'Usted no entreno el modelo de textura con los vectores de textura'
assert shape_model4.n_features_in_==len(descript_shape_train[0]), f'Usted no entreno el modelo de forma con los vectores de forma'
assert color_model4.n_features_in_==len(descript_color_train[0]), f'Usted no entreno el modelo de color con los vectores de color'
```

### Parte 3: Random Forest

Inicialice 4 modelos para cada descriptor (12 modelos) variando el valor de n_estimators y max_features. Entrene sus respectivos modelos.

<!-- grade_id: cell-bf477d7d1b33e1af, points: 0, grade: false, solution: true -->
```python
from sklearn.ensemble import RandomForestClassifier

n_estimators1='' #valor 1 para n_estimators
n_estimators2='' #valor 2 para n_estimators

max_features1='' #max_features 1
max_features2='' #max_features 2

color_model1_RF='' #max_features 1 y n_estimators 1
color_model2_RF='' #max_features 1 y n_estimators 2
color_model3_RF='' #max_features 2 y n_estimators 1
color_model4_RF='' #max_features 2 y n_estimators 2

shape_model1_RF='' #max_features 1 y n_estimators 1
shape_model2_RF='' #max_features 1 y n_estimators 2
shape_model3_RF='' #max_features 2 y n_estimators 1
shape_model4_RF='' #max_features 2 y n_estimators 2

texture_model1_RF='' #max_features 1 y n_estimators 1
texture_model2_RF='' #max_features 1 y n_estimators 2
texture_model3_RF='' #max_features 2 y n_estimators 1
texture_model4_RF='' #max_features 2 y n_estimators 2

# YOUR CODE HERE

n_estimators1 = 10
n_estimators2 = 100

max_features1 = 'sqrt' 
max_features2 = 'log2'

color_model1_RF = RandomForestClassifier(n_estimators=n_estimators1, max_features=max_features1).fit(descript_color_train, labels_color_train)
color_model2_RF = RandomForestClassifier(n_estimators=n_estimators2, max_features=max_features1).fit(descript_color_train, labels_color_train)
color_model3_RF = RandomForestClassifier(n_estimators=n_estimators1, max_features=max_features2).fit(descript_color_train, labels_color_train)
color_model4_RF = RandomForestClassifier(n_estimators=n_estimators2, max_features=max_features2).fit(descript_color_train, labels_color_train)

shape_model1_RF = RandomForestClassifier(n_estimators=n_estimators1, max_features=max_features1).fit(descript_shape_train, labels_shape_train)
shape_model2_RF = RandomForestClassifier(n_estimators=n_estimators2, max_features=max_features1).fit(descript_shape_train, labels_shape_train)
shape_model3_RF = RandomForestClassifier(n_estimators=n_estimators1, max_features=max_features2).fit(descript_shape_train, labels_shape_train)
shape_model4_RF = RandomForestClassifier(n_estimators=n_estimators2, max_features=max_features2).fit(descript_shape_train, labels_shape_train)

texture_model1_RF = RandomForestClassifier(n_estimators=n_estimators1, max_features=max_features1).fit(descript_texture_train, labels_texture_train)
texture_model2_RF = RandomForestClassifier(n_estimators=n_estimators2, max_features=max_features1).fit(descript_texture_train, labels_texture_train)
texture_model3_RF = RandomForestClassifier(n_estimators=n_estimators1, max_features=max_features2).fit(descript_texture_train, labels_texture_train)
texture_model4_RF = RandomForestClassifier(n_estimators=n_estimators2, max_features=max_features2).fit(descript_texture_train, labels_texture_train)
```

<!-- grade_id: cell-bc5291dd7c4ed639, points: 0.7, grade: true, solution: false -->
```python
model1_params_color_RF=color_model1_RF.get_params()
model2_params_color_RF=color_model2_RF.get_params()
model3_params_color_RF=color_model3_RF.get_params()
model4_params_color_RF=color_model4_RF.get_params()

model1_params_shape_RF=shape_model1_RF.get_params()
model2_params_shape_RF=shape_model2_RF.get_params()
model3_params_shape_RF=shape_model3_RF.get_params()
model4_params_shape_RF=shape_model4_RF.get_params()

model1_params_texture_RF=texture_model1_RF.get_params()
model2_params_texture_RF=texture_model2_RF.get_params()
model3_params_texture_RF=texture_model3_RF.get_params()
model4_params_texture_RF=texture_model4_RF.get_params()

assert model1_params_color_RF['n_estimators']==n_estimators1, f'El valor de n_estimators deberia ser {n_estimators1} y es {model1_params_color_RF["n_estimators"]}'
assert model2_params_shape_RF['n_estimators']==n_estimators2, f'El valor de n_estimators deberia ser {n_estimators2} y es {model2_params_shape_RF["n_estimators"]}'
assert model3_params_texture_RF['max_features']==max_features2, f'El max_features deberia ser {max_features2} y es {model3_params_texture_RF["max_features"]}'
assert texture_model4_RF.n_features_in_==len(descript_texture_train[0]), f'Usted no entreno el modelo de textura con los vectores de textura'
assert shape_model4_RF.n_features_in_==len(descript_shape_train[0]), f'Usted no entreno el modelo de forma con los vectores de forma'
assert color_model4_RF.n_features_in_==len(descript_color_train[0]), f'Usted no entreno el modelo de color con los vectores de color'
```

### Parte 4: MLP

Por ultimo, usaremos redes neuronales como un clasificador, el MLP. Las redes neuronales reciben como entrada un vector de caracteristicas y, mediante multiplicaciones matriciales, funciones de activación y el descenso del gradiente, las redes neuronales aprenden a clasificar los descriptores optimizando una función de perdida. Inicialice 4 modelos para cada descriptor (12 modelos) variando el valor de hidden_layer_sizes y learning_rate_init. Entrene sus respectivos modelos.

<!-- grade_id: cell-565641e8f7da3569, points: 0, grade: false, solution: true -->
```python
import numpy as np
from sklearn.neural_network import MLPClassifier

hidden_layer_sizes1='' #valor 1 para hidden_layer_sizes
hidden_layer_sizes2='' #valor 2 para hidden_layer_sizes

learning_rate_init1='' #learning_rate_init 1
learning_rate_init2='' #learning_rate_init 2

color_model1_MLP='' #learning_rate_init 1 y hidden_layer_sizes1
color_model2_MLP='' #learning_rate_init 1 y hidden_layer_sizes2
color_model3_MLP='' #learning_rate_init 2 y hidden_layer_sizes1
color_model4_MLP='' #learning_rate_init 2 y hidden_layer_sizes2

shape_model1_MLP='' #learning_rate_init 1 y hidden_layer_sizes1
shape_model2_MLP='' #learning_rate_init 1 y hidden_layer_sizes2
shape_model3_MLP='' #learning_rate_init 2 y hidden_layer_sizes1
shape_model4_MLP='' #learning_rate_init 2 y hidden_layer_sizes2

texture_model1_MLP='' #learning_rate_init 1 y hidden_layer_sizes1
texture_model2_MLP='' #learning_rate_init 1 y hidden_layer_sizes2
texture_model3_MLP='' #learning_rate_init 2 y hidden_layer_sizes1
texture_model4_MLP='' #learning_rate_init 2 y hidden_layer_sizes2

# YOUR CODE HERE

hidden_layer_sizes1 = (50,)
hidden_layer_sizes2 = (50, 50)

learning_rate_init1 = 0.0001
learning_rate_init2 = 0.001

# subir max_iter?

color_model1_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes1, learning_rate_init=learning_rate_init1).fit(descript_color_train, labels_color_train)
color_model2_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes2, learning_rate_init=learning_rate_init1).fit(descript_color_train, labels_color_train)
color_model3_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes1, learning_rate_init=learning_rate_init2).fit(descript_color_train, labels_color_train)
color_model4_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes2, learning_rate_init=learning_rate_init2).fit(descript_color_train, labels_color_train)

shape_model1_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes1, learning_rate_init=learning_rate_init1).fit(descript_shape_train, labels_shape_train)
shape_model2_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes2, learning_rate_init=learning_rate_init1).fit(descript_shape_train, labels_shape_train)
shape_model3_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes1, learning_rate_init=learning_rate_init2).fit(descript_shape_train, labels_shape_train)
shape_model4_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes2, learning_rate_init=learning_rate_init2).fit(descript_shape_train, labels_shape_train)

texture_model1_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes1, learning_rate_init=learning_rate_init1).fit(descript_texture_train, labels_texture_train)
texture_model2_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes2, learning_rate_init=learning_rate_init1).fit(descript_texture_train, labels_texture_train)
texture_model3_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes1, learning_rate_init=learning_rate_init2).fit(descript_texture_train, labels_texture_train)
texture_model4_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes2, learning_rate_init=learning_rate_init2).fit(descript_texture_train, labels_texture_train)

```

Output:
```
C:\Users\eduar\anaconda3\Lib\site-packages\sklearn\neural_network\_multilayer_perceptron.py:691: ConvergenceWarning: Stochastic Optimizer: Maximum iterations (200) reached and the optimization hasn't converged yet.
  warnings.warn(
C:\Users\eduar\anaconda3\Lib\site-packages\sklearn\neural_network\_multilayer_perceptron.py:691: ConvergenceWarning: Stochastic Optimizer: Maximum iterations (200) reached and the optimization hasn't converged yet.
  warnings.warn(
C:\Users\eduar\anaconda3\Lib\site-packages\sklearn\neural_network\_multilayer_perceptron.py:691: ConvergenceWarning: Stochastic Optimizer: Maximum iterations (200) reached and the optimization hasn't converged yet.
  warnings.warn(
C:\Users\eduar\anaconda3\Lib\site-packages\sklearn\neural_network\_multilayer_perceptron.py:691: ConvergenceWarning: Stochastic Optimizer: Maximum iterations (200) reached and the optimization hasn't converged yet.
  warnings.warn(
C:\Users\eduar\anaconda3\Lib\site-packages\sklearn\neural_network\_multi
```

```python
# pruebas individuales de convergencia:

hidden_layer_sizesP = (50,)
learning_rate_initP = 0.001

color_model1_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes1, learning_rate_init=learning_rate_init1).fit(descript_color_train, labels_color_train)

shape_model1_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes1, learning_rate_init=learning_rate_init1).fit(descript_shape_train, labels_shape_train)

texture_model1_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes1, learning_rate_init=learning_rate_init1).fit(descript_texture_train, labels_texture_

```

<!-- grade_id: cell-bf5a3c9732d027fa, points: 0.7, grade: true, solution: false -->
```python
model1_params_color_MLP=color_model1_MLP.get_params()
model2_params_color_MLP=color_model2_MLP.get_params()
model3_params_color_MLP=color_model3_MLP.get_params()
model4_params_color_MLP=color_model4_MLP.get_params()

model1_params_shape_MLP=shape_model1_MLP.get_params()
model2_params_shape_MLP=shape_model2_MLP.get_params()
model3_params_shape_MLP=shape_model3_MLP.get_params()
model4_params_shape_MLP=shape_model4_MLP.get_params()

model1_params_texture_MLP=texture_model1_MLP.get_params()
model2_params_texture_MLP=texture_model2_MLP.get_params()
model3_params_texture_MLP=texture_model3_MLP.get_params()
model4_params_texture_MLP=texture_model4_MLP.get_params()

assert model1_params_color_MLP['hidden_layer_sizes']==hidden_layer_sizes1, f'El valor de hidden_layer_sizes1 deberia ser {hidden_layer_sizes1} y es {model1_params_color_MLP["hidden_layer_sizes"]}'
assert model2_params_shape_MLP['hidden_layer_sizes']==hidden_layer_sizes2, f'El valor de hidden_layer_sizes1 deberia ser {hidden_layer_sizes2} y es {model2_params_shape_MLP["hidden_layer_sizes"]}'
assert model3_params_texture_MLP['learning_rate_init']==learning_rate_init2, f'El learning_rate_init deberia ser {learning_rate_init2} y es {model3_params_texture_MLP["learning_rate_init"]}'
assert texture_model4_MLP.n_features_in_==len(descript_texture_train[0]), f'Usted no entreno el modelo de textura con los vectores de textura'
assert shape_model4_MLP.n_features_in_==len(descript_shape_train[0]), f'Usted no entreno el modelo de forma con los vectores de forma'
assert color_model4_MLP.n_features_in_==len(descript_color_train[0]), f'Usted no entreno el modelo de color con los vectores de color'
```

### Parte 5: Validación

Una vez entrenados nuestros modelos, debemos validarlos para obtener el mejor de ellos, tanto en descriptor como en clasificador. Haga uso de el descriptor de sus imagenes de validación y las metricas de precisión, cobertura y F medida para evaluar todos sus modelos. Recuerde que debe validar los modelos con su respectivo descriptor (color con color, textura con textura y forma con forma)

<!-- grade_id: cell-433ceab2a37de015, points: 0, grade: false, solution: true -->
```python
from sklearn.metrics import f1_score, precision_score, recall_score

# YOUR CODE HERE

f1_manual = lambda p, c: 2*p*c/(p+c)

def metricas(annots, predictions, carpeta='None', nombre_modelo='None', manual = False):
    
    precision = precision_score(annots, predictions, average='weighted')
    cobertura = recall_score(annots, predictions, average='weighted')
    f1 = f1_score(annots, predictions, average='weighted') if not manual else f1_manual(precision, cobertura)
        
    print()
    print(f'El f1 de {carpeta} del modelo {nombre_modelo} es {f1}')
    print(f'La precision de {carpeta} del modelo {nombre_modelo} es {precision}')
    print(f'La cobertura de {carpeta} del modelo {nombre_modelo} es {cobertura}')
    
def nombrar_modelo(modelo):
    name = ''
    if isinstance(modelo, SVC):
        name = f'SVC con C={modelo.C} y kernel={modelo.kernel}'
    elif isinstance(modelo, RandomForestClassifier):
        name = f'RandomForest con n_estimators={modelo.n_estimators} y max_features={modelo.max_features}'
    elif isinstance(modelo, MLPClassifier):
        name = f'MLP con hidden_layer_sizes={modelo.hidden_layer_sizes} y learning_rate_init={modelo.learning_rate_init}'
    else:
        raise Exception('Tipo de modelo inválido')
    return name    
    
def predict_classifier(models, descript, labels, carpeta='None'):
    for model in models:
        predictions = model.predict(descript)
        metricas(labels, predictions, carpeta=carpeta, nombre_modelo=nombrar_modelo(model))
 
# el mejor fue el de 17 textones (exp. 4)
descript_texture_valid, labels_texture_valid = load_descriptors_labels('valid_texture_4.npy')
print('El mejor modelo de textura encontrado en la entrega 2 fue el de 17 textones.')

# el mejor fue el de joint+lab (exp. 3)
descript_color_valid, labels_color_valid = load_descriptors_labels('valid_color_3.npy')
print('El mejor modelo de color encontrado en la entrega 2 fue el de joint+lab.')

# el mejor fue el de 14 orientaciones y 4 pixeles_por_celda (exp. 1)
descript_shape_valid, labels_shape_valid = load_descriptors_labels('valid_shape_1.npy')
print('El mejor modelo de forma encontrado en la entrega 2 fue el de 14 orientaciones y 4 pixeles por celda.')

# las .npy se envían en la entrega? SI
# UndefinedMetricWarning: Precision is ill-defined and being set to 0.0 in labels with no predicted samples. Use `zero_division` parameter to control this behavior.
# COVERGENCE  
```

Output:
```
El mejor modelo de textura encontrado en la entrega 2 fue el de 17 textones.
El mejor modelo de color encontrado en la entrega 2 fue el de joint+lab.
El mejor modelo de forma encontrado en la entrega 2 fue el de 14 orientaciones y 4 pixeles por celda.
```

```python
models_color = [color_model1, color_model2, color_model3, color_model4, 
                color_model1_RF, color_model2_RF, color_model3_RF, color_model4_RF,
                color_model1_MLP, color_model2_MLP, color_model3_MLP, color_model4_MLP]

models_shape = [shape_model1, shape_model2, shape_model3, shape_model4,
                shape_model1_RF, shape_model2_RF, shape_model3_RF, shape_model4_RF,
                shape_model1_MLP, shape_model2_MLP, shape_model3_MLP, shape_model4_MLP]

models_texture = [texture_model1, texture_model2, texture_model3, texture_model4, 
                  texture_model1_RF, texture_model2_RF, texture_model3_RF, texture_model4_RF,
                  texture_model1_MLP, texture_model2_MLP, texture_model3_MLP, texture_model4_MLP]

print('Probamos los modelos de COLOR:')
predict_classifier(models_color, descript_color_valid, labels_color_valid, 'valid')

print('\n\nProbamos los modelos de FORMA:')
predict_classifier(models_shape, descript_shape_valid, labels_shape_valid, 'valid')

print('\n\nProbamos los modelos de TEXTURA:')
predict_classifier(models_texture, descript_texture_valid, labels_texture_valid, 'valid')
```

Output:
```
Probamos los modelos de COLOR:

El f1 de valid del modelo SVC con C=0.1 y kernel=linear es 0.3122619962615379
La precision de valid del modelo SVC con C=0.1 y kernel=linear es 0.31641598598120335
La cobertura de valid del modelo SVC con C=0.1 y kernel=linear es 0.336

El f1 de valid del modelo SVC con C=10 y kernel=linear es 0.33153887998415765
La precision de valid del modelo SVC con C=10 y kernel=linear es 0.32371473685412766
La cobertura de valid del modelo SVC con C=10 y kernel=linear es 0.356

El f1 de valid del modelo SVC con C=0.1 y kernel=rbf es 0.29191327698603964
La precision de valid del modelo SVC con C=0.1 y kernel=rbf es 0.3066807753649859
La cobertura de valid del modelo SVC con C=0.1 y kernel=rbf es 0.332

El f1 de valid del modelo SVC con C=10 y kernel=rbf es 0.4177096163846356
La precision de valid del modelo SVC con C=10 y kernel=rbf es 0.42170364322315446
La cobertura de valid del modelo SVC con C=10 y kernel=rbf es 0.424

El f1 de valid del modelo Random Forest con
```

```python
"""
color_model1='' #kernel 1 y C1
color_model2='' #kernel 1 y C2
color_model3='' #kernel 2 y C1
color_model4='' #kernel 2 y C2

shape_model1='' #kernel 1 y C1
shape_model2='' #kernel 1 y C2
shape_model3='' #kernel 2 y C1
shape_model4='' #kernel 2 y C2

texture_model1='' #kernel 1 y C1
texture_model2='' #kernel 1 y C2
texture_model3='' #kernel 2 y C1
texture_model4='' #kernel 2 y C2


color_model1_RF='' #max_features 1 y n_estimators 1
color_model2_RF='' #max_features 1 y n_estimators 2
color_model3_RF='' #max_features 2 y n_estimators 1
color_model4_RF='' #max_features 2 y n_estimators 2

shape_model1_RF='' #max_features 1 y n_estimators 1
shape_model2_RF='' #max_features 1 y n_estimators 2
shape_model3_RF='' #max_features 2 y n_estimators 1
shape_model4_RF='' #max_features 2 y n_estimators 2

texture_model1_RF='' #max_features 1 y n_estimators 1
texture_model2_RF='' #max_features 1 y n_estimators 2
texture_model3_RF='' #max_features 2 y n_estimators 1
texture_model4_RF='' #max_features 2 y n_estimators 2


color_model1_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes1, learning_rate_init=learning_rate_init1).fit(descript_color_train, labels_color_train)
color_model2_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes2, learning_rate_init=learning_rate_init1).fit(descript_color_train, labels_color_train)
color_model3_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes1, learning_rate_init=learning_rate_init2).fit(descript_color_train, labels_color_train)
color_model4_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes2, learning_rate_init=learning_rate_init2).fit(descript_color_train, labels_color_train)

shape_model1_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes1, learning_rate_init=learning_rate_init1).fit(descript_shape_train, labels_shape_train)
shape_model2_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes2, learning_rate_init=learning_rate_init1).fit(descript_shape_train, labels_shape_train)
shape_model3_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes1, learning_rate_init=learning_rate_init2).fit(descript_shape_train, labels_shape_train)
shape_model4_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes2, learning_rate_init=learning_rate_init2).fit(descript_shape_train, labels_shape_train)

texture_model1_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes1, learning_rate_init=learning_rate_init1).fit(descript_texture_train, labels_texture_train)
texture_model2_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes2, learning_rate_init=learning_rate_init1).fit(descript_texture_train, labels_texture_train)
texture_model3_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes1, learning_rate_init=learning_rate_init2).fit(descript_texture_train, labels_texture_train)
texture_model4_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer_sizes2, learning_rate_init=learning_rate_init2).fit(descript_texture_train, labels_texture_train)
"""

print()
```

```python
# Experimentación con MLP para evitar warning de convergencia

# YOUR CODE HERE

hidden_layer_sizes1 = (50,)
hidden_layer_sizes2 = (50, 50)

learning_rate_init1 = 0.0001
learning_rate_init2 = 0.001

hidden_layers = [(50,), (50, 50), (50, 50, 50), (50, 50, 50, 50), (100,), (100, 100), (100, 100, 100), (100, 100, 100, 100)]

learning_rates = [0.0001, 0.001, 0.01, 0.1]

for hidden_layer in hidden_layers:
    for learning_rate in learning_rates:
        print(f'hidden_layer_size: {hidden_layer}, learning_rate:, {learning_rate}')
        print('color')
        color_model_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer, learning_rate_init=learning_rate).fit(descript_color_train, labels_color_train)
        print('shape')
        shape_model_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer, learning_rate_init=learning_rate).fit(descript_shape_train, labels_shape_train)
        print('texture')
        texture_model_MLP = MLPClassifier(hidden_layer_sizes=hidden_layer, learning_rate_init=learning_rate).fit(descript_texture_train, labels_texture_train)
```

## Articulo

Una vez realizado el procedimiento es fundamental que se continue con la construcción del articulo de forma progresiva, por lo cual incluiremos lo que hicimos en esta entrega.

### Experimentos

#### Clasificadores
En la entrega pasada hemos experimentado con los descriptores unicamente. En esta ocasión, experimentamos con los parametros de diferentes clasificadores. Agregue los resultados de su experimentación usando los diferentes clasificadores junto con una breve descripción de sus resultados (no discusión). Puede guiarse respondiendo las siguientes preguntas:

- ¿Cuál fue su mejor método?
- ¿Hubo algún parámetro que influyera mas? (i.e. Un C mas grande en el SVM dió resultados mejores en la clasificaición que modificar el kernel.)
- ¿Qué combinación no funcionó de forma correcta?
- ¿Existía alguna tendencia?
- ¿Algun clasificador tiene un desempeño mejor que el baseline, alguno tiene un desempeño peor?
- ¿Cual clasificador funcionó mejor que el resto?

### Discusión

#### Clasificadores
Incluyan un análisis de lo obtenido en sus experimentos. Tenga en cuenta resultados cuantitativos y cualitativos. Teorice sobre las mejoras o las desmejoras de los modelos. Puede guiarse respondiendo las siguientes preguntas:

- ¿Algun clasificador no mejoro el baseline? ¿cual puede ser la razon?
- ¿Cuales son los sustentos teoricos de los resultados? ¿Hay algun resultado anti-intuitivo?
- ¿Los resultados de los descriptores son consistentes en todos los clasificadores? ¿por que?

#### Metodo final
Por ultimo, debe determinar cual es su mejor modelo de clasificación y su mejor descriptor. Haga un analisis en profundidad de la combinación de parametros y del descriptor final utilizado. Describa tambien las caracteristicas de sus resultados, tiene una alta precisión o una alta cobertura y las posibles implicaciones de esto en la vida real. 

### Conclusiones
Concluya acerca de los resultados obtenidos en esta base de datos y lo novedoso de su aproximación. ¿Cree que su aproximación es suficiente? Como mejoraria su modelo para obtener mejores resultados. 

```python
from utils import *

converter("Entrega 3")
```