# MiniProyecto 2: Segmentación de higado en volumenes MRI

## Entrega 1: 

En el anterior miniproyecto empezamos a relacionarnos con las tareas de la visión por computador y la metodología de investigación, que va desde la visualización y el entendimiento del problema a resolver, pasando por la planeación de experimentos y el pre-procesamiento, para finalizar con el desarrollo de un algoritmo completo y su metodología de evaluación. En el anterior miniproyecto nos centramos en la detección, una tarea de clasificación binaria donde buscamos remarcar o delimitar un objeto en una ventana. Ahora, nos familiarizaremos con la tarea de segmentación, con la que ya pudimos experimentar en la entrega pasada, pero ahora siendo el problema principal y no un vehiculo para la detección. Debemos recordar que la segmentación es una refinación de la tarea de clasificación, donde queremos asignarle a cada pixel una etiqueta, en este caso será una clasificación binaria donde la clase 0 es el fondo y la clase 1 será el higado. 

### Parte 1: Exploración de base de datos

Como aprendimos anteriormente, lo primero que debemos hacer es familiarizarnos con las imágenes de la base de datos. En esta ocasión no tenemos imagenes como tal sino volumenes ya que tratamos con imagenes medicas, en concreto, MRI.

#### Parte 1.1: Importación de imagenes y anotaciones

Para poder abrir los archivos de MRI debemos usar la libreria **nibabel**. Use la siguiente celda para cargar el volumen original.nii y la segmentación segmentation.nii de la carpeta E1. 

> - **Nota** Para poder abrir los volumenes debe usar la función *load* y la función *get_fdata*. Puede ver un ejemplo en este link https://nipy.org/nibabel/gettingstarted.html. 

```python
#!pip install nibabel
```

<!-- grade_id: cell-06681aeb7ffa05c8, points: 0, grade: false, solution: true -->
```python
original_vol='' # Volumen original original.nii
seg_vol='' #Volumen de segmentación segmentation.nii

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d9fb5601ccedf7aa, points: 0.1, grade: true, solution: false -->
```python
assert np.isclose(np.sum(original_vol),5004343,0.01), f'El volumen original.nii no es correcto'
assert np.isclose(np.sum(seg_vol),555218,0.01), f'El volumen de segmentación no es el correcto'
```

#### Parte 1.2 Distribución de la imagen y segmentación
Ahora vamos a explorar la distribución de las clases en la segmentación. Al igual que en los problemas de detección, tendremos una distribución de etiquetas dentro de cada dato, recuerden que nuestros datos son volumenes. Use la siguiente celda para encontrar la cantidad de pixeles de cada clase en la volumen original. Tambien saque los porcentajes de estas clases. Todo el dataset tiene 3 labels y cada label es un numero, tengalo en cuenta.

> **Recomendación:** Use np.unique (busque la documentación) para encontrar los labels de las segmentaciones, sabemos que son numeros pero no que numeros son.

<!-- grade_id: cell-219425899d701541, points: 0, grade: false, solution: true -->
```python
labels='' #arreglo con las etiquetas de la base de datos de train

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-48319ebf83b2f0d0, points: 0.1, grade: true, solution: false -->
```python
assert labels.shape[0] == 3, f'El número de etiquetas debería ser 3'
assert np.isclose(np.sum(labels), 3, 0.1), f'Obtuvieron los labels incorrectamente'
```

<!-- grade_id: cell-0e170d22341fd2d8, points: 0, grade: false, solution: true -->
```python
label_0='' #cantidad de voxeles con la etiqueta 0
label_1='' #cantidad de voxeles con la etiqueta 1
label_2='' #cantidad de voxeles con la etiqueta 2

# YOUR CODE HERE
raise NotImplementedError()
```

```python
print ("{:<15} {:<20}".format('Clase','Total voxeles por clase'))
print ('----------------------------------------------------------------------------')
print ("{:<15} {:<20}".format(labels[0],label_0))
print ("{:<15} {:<20}".format(labels[1],label_1))
print ("{:<15} {:<20}".format(labels[2],label_2))
```

<!-- grade_id: cell-2a8f1b940c0ae82b, points: 0.1, grade: true, solution: false -->
```python
assert np.isclose(np.sum(label_0), 19108173, 0.01), f'El número de píxeles en la clase 0 debería ser 19108173 y no {np.sum(label_0)}'
assert np.isclose(np.sum(label_1), 550036, 0.01), f'El número de píxeles en la clase 1 debería ser 19108173 y no {np.sum(label_1)}'
assert np.isclose(np.sum(label_1), 550036, 0.01), f'El número de píxeles en la clase 1 debería ser 19108173 y no {np.sum(label_1)}'
```

#### Parte 1.3: Visualización de anotaciones
Veamos ahora los datos con los que vamos a trabajar. Sabemos que nuestros datos originales son MRI, son volumenes, pero aun no sabemos como serán nuestras anotaciones. Esto es importante porque, como todo problema de vision por computador, debemos saber cuales son nuestros inputs y cuales queremos que sean nuestros outputs. En la siguiente celda use el volumen original.nii y busque cual es el slice de la anotación (volumen segmentado) que más area de segmentación tiene en el eje axial. Haga un plot de del corte y de la segmentacion más grande. 

> **Recuerde:** Tenemos el label del fondo (0) y las demas labels, pero nosotros solo queremos tratar con una segmentación binaria. De ahora en adelante, las labels que son diferentes de 0 serán 1.

<!-- grade_id: cell-935bb3352f446a7e, points: 0, grade: false, solution: true -->
```python
slice_axial='' #Slice con mayor area de higado en el eje axial
area_axial='' #Mayor area de higado en eje axial
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e37b7df59d8eb0e6, points: 0.1, grade: true, solution: false -->
```python
assert len(np.unique(seg_vol))==2, f'Las etiquetas del volumen de anotacion deben ser 2 y son {len(np.unique(segmen1))}, puede que haya guardado el volumen 1 en esta variable o no haya contado la etiqueta 1 y 2 como 1.'
assert np.isclose(slice_axial,65,0.01), 'El eje mayor no es el encontrado. El eje axial es la ultima coordenada, la transversal es la segunda y la sagital es la primera'
assert np.isclose(np.sum(seg_vol[:,:,slice_axial]>0)-area_axial,0), 'El area del higado guardado en la variable no es igual al slice de maxima area determinada'
```

### Parte 2: Anotación
El problema que queremos tratar es el de la segmentación del higado. La segmentación que realizaremos se basará directamente en las intensidades de nuestra imagen. Un punto a destacar es la diferencia que tenemos con la anterior base de datos, ahora no contamos con imagenes a color y muchos tegidos tendrán una intensidad similar a la de nuestro organo de interes.

#### Parte 2.1: Histograma
Usted cuenta con volumenes de anotación, si visualizó con cuidado sus anotaciones sabrá que son volumenes de etiquetas (0,1,2 aunque consideramos el 2 y el 1 como una misma etiqueta) Utilice este volumen de anotación para extraer el higado con sus intensidades originales. Posteriormente, realice un plot mostrando el organo y su histograma (no nos importan las intensidades de 0 porque son el fondo). Utilice el volumen original.nii y la segmentación segmentation.nii de la carpeta E1. Debe hacer el histograma para todo el higado, no solo una slice.

> - **Pista:** Si su volumen de segmentación tiene valores de 0 o 1, si se multiplica cualquier valor por 0, este desaparecerá, y todo aquel voxel que pertenezca a la segmentación será multiplicado por 1 y conservará su intensidad original. Haga el histograma sin tener en cuenta el fondo.

<!-- grade_id: cell-802da64e7fab52b2, points: 0, grade: false, solution: true -->
```python
Higado='' #volumen que contendrá solo el higado de la MRI
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3a10f074b299c437, points: 0.2, grade: true, solution: false -->
```python
assert np.max(Higado)<=1.1, 'Seguramente su segmentación tiene la etiqueta 2'
assert np.isclose(np.mean(Higado),0.025,0.1), f'El valor medio del organo es {np.mean(Higado)} y no deberia ser este, quizá esta normalizando o no está tomando todo el organo'
```

#### Parte 2.2: Ruido
Un punto importante antes de realizar nuestro modelo de segmentación es el estado de las imagenes. Como utilizaremos las intesidades para realizar nuestra segmentación debemos asegurarnos que no tenemos ruido que afecte nuestras intensidades originales. 

Cargue los volumenes corruptos que se encuentran en la carpeta E1/corrupted_volumes y haga el mismo procedimeinto anterior. Haga un plot del higado segmentado corrupto y su respectivo histograma.

<!-- grade_id: cell-e51009b1d437d2bc, points: 0, grade: false, solution: true -->
```python
salt_peper='' #volumen con ruido de sal y pimienta
gaussian='' #volumen con ruido gaussiano
Higado_salt_peper='' #volumen que contendrá solo el higado de la MRI con ruido sal y pimienta
Higado_gaussian='' #volumen que contendrá solo el higado de la MRI con ruido gaussiano

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2e3e7c96b9bfd8a4, points: 0.3, grade: true, solution: false -->
```python
assert np.max(Higado_salt_peper)<=1.1, 'Seguramente su segmentación tiene la etiqueta 2'
assert np.isclose(np.sum(Higado_salt_peper!=0),497413,0.01), 'El volumen total del higado es incorrecto, no lo esta tomando todo'
assert np.isclose(np.mean(Higado_salt_peper),0.023,0.1), f'El valor medio del organo es {np.mean(Higado)} y no deberia ser este, quizá esta normalizando o no está tomando todo el organo'

assert np.max(Higado_gaussian)<=1.4, 'Seguramente su segmentación tiene la etiqueta 2'
assert np.isclose(np.sum(Higado_gaussian!=0),552627,0.01), 'El volumen total del higado es incorrecto, no lo esta tomando todo'
assert np.isclose(np.mean(Higado_gaussian),0.025,0.1), f'El valor medio del organo es {np.mean(Higado)} y no deberia ser este, quizá esta normalizando o no está tomando todo el organo'
```

### Parte 3: Remoción de ruido
Una parte importante a la hora de resolver problemas de vision por computador es el preprocesamiento. El ruido puede tener diversas fuentes y puede corromper la imagen de diferentes formas. Es nuestro trabajo manipular las imagenes de tal forma que sean más faciles de procesar por nuestro algoritmo de predicción. Es tambien muy importante hacer un diseño exhaustivo del procedimiento de filtrado que realizaremos, ya que no todos los filtros nos ayudaran a el procesamiento de las imagenes. Algunos filtros son muy especializados en el tipo de ruido que logrará filtrar, algunos filtros pueden magnificar el ruido o incluso dañar la imagen.

#### Parte 3.1: Filtrado lineal

Una primera idea puede ser el uso de la cross correlacion para el trabajo del filtrado. Use la función de scipy.signal.correlate y diseñe un filtro medio 3x3. Tambien use la función scipy.ndimage.gaussian_filter para hacer un filtro gaussiano y experimente con el sigma para obtener mejores resultados. Lea los dos volumenes de E1/corrupted_volumes, extraiga el slice axial encontrado en el punto 1 y use estos filtros, guarde los resultados en variables. Haga un subplot de 4 filas y 3 columnas, en las filas iran los diferentes ruidos con los diferentes filtros (ruido gaussiano con filtro medio, ruido gaussiano con filtro gaussiano, etc.) y en las columnas la imagen original (solo el higado), la imagen procesada y el histograma de la imagen procesada.

> - **Nota:** Procese las imagenes orignales y luego si extraiga unicamente el higado, no lo haga al reves. Use el parametro "same" para mantener el tamaño de la imagen original.

<!-- grade_id: cell-73f9d0e6f35329b3, points: 0, grade: false, solution: true -->
```python
salt_peper_axial='' #slice del volumen con ruido sal y pimienta
gaussian_axial='' #slice del volumen con ruido gaussiano

media_filter='' #arreglo ndarray del filtro medio

media_salt_peper='' #slice de ruido sal y pimienta filtrado con filtro medio
gaussian_salt_peper='' #slice de ruido sal y pimienta filtrado con filtro gaussiano, sigma 3
media_gaussian='' #slice de ruido gaussiano filtrado con filtro medio
gaussian_gaussian='' #slice de ruido gaussiano filtrado con filtro gaussiano, sigma 3

gaussian_salt_peper_higado='' #slice de ruido sal y pimienta filtrado con filtro gaussiano solo del higado
gaussian_gaussian_higado='' #slice de ruido gaussiano filtrado con filtro gaussiano solo del higado
media_salt_peper_higado='' #slice de ruido sal y pimienta filtrado con filtro medio solo del higado
media_gaussian='' #slice de ruido gaussiano filtrado con filtro medio solo del higado

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-03eb68701d71bcc4, points: 0.1, grade: true, solution: false -->
```python
assert np.sum(media_filter)==1, 'La suma de todos los elementos del filtro deben sumar 1'
assert np.sum(media_filter==1/9)==9, 'Su filtro medio esta mal construido'
x,y=gaussian_salt_peper_higado.shape
x2,y2=salt_peper_axial.shape
assert x==x2 and y==y2, 'La imagen filtrada debe ser del mismo tamaño de la imagen original'
```

#### Parte 3.2: Filtrado no lineal
El filtrado lineal nos es muy util para obtener caracteristicas visuales de las imagenes y tambien para reducir cierto tipo de ruido. Sin embargo, pudieron observar cualitativamente las limitaciones de este tipo de filtrado. Una idea interesante puede ser el uso de filtrado no lineal para la eliminación de ruido y menor destrucción de la imagen original. Use los filtros maximo, minimo, mediano y mediano adaptativo para remover el ruido de las imagenes del anterior punto (puede usar la implementación del tutorial) haga una función que pueda hacer todos estos filtros segun se indique por parametro, tome en cuenta la condición de frontera como 0s. Haga un subplot de 2 filas 5 columnas, donde las filas son los diferentes ruidos y las columnas son la imagen original y los 4 tipos de filtros.

<!-- grade_id: cell-f44b24c9bc8759d3, points: 0, grade: false, solution: true -->
```python
def No_linear_filter(Im,filter_type,window,max_window=3):
    '''
    :param Im: Imagen oringinal.
    :param filter_type: parametro que determina el tipo de filtrado a usar. 
                    Puede tomar de valores: 'min','max', 'median' or 'median_adapt'
    :param window: Tamaño de la ventana de analisis.
    :param max_window: Maximo tamaño de la ventana si se esta haciendo un filtro mediano adaptativo.
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
    return new_image
```

```python
Max_salt_peper=No_linear_filter(salt_peper_axial,'max',3)
Min_salt_peper=No_linear_filter(salt_peper_axial,'min',3)
Median_salt_peper=No_linear_filter(salt_peper_axial,'median',3)
MedianAdapt_salt_peper=No_linear_filter(salt_peper_axial,'median_adapt',3,11)

Max_gaussian=No_linear_filter(gaussian_axial,'max',3)
Min_gaussian=No_linear_filter(gaussian_axial,'min',3)
Median_gaussian=No_linear_filter(gaussian_axial,'median',3)
MedianAdapt_gaussian=No_linear_filter(gaussian_axial,'median_adapt',3,11)
```

#### Parte 3.3 Evaluación cuantitativa

Hemos podido experimentar con los diferentes filtros que tenemos a nuestra disposición obteniendo resultados mejores o peores. La comparación que hemos podido tener es meramente cualitativa, sin embargo, contamos con una herramienta que nos permite saber que tan eficaz ha sido nuestro filtrado, el PSNR. Cree una función que entre como parametro 2 imagenes de un solo canal y retorne su PSNR. Considere el maximo de la imagen como 1.

<!-- grade_id: cell-14e1dc8dbf5bfc95, points: 0, grade: false, solution: true -->
```python
original='' #slice del volumen original sin ruido
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-43380e2b29a21a51, points: 0, grade: false, solution: true -->
```python
def PSNR(Im,Filtered):
    '''
    :param Im: Imagen oringinal.
    :param Filtered: Imagen filtrada.
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
    return result
```

```python
Value_salt_peper_media=PSNR(original,media_salt_peper)
Value_salt_peper_gaussian=PSNR(original,gaussian_salt_peper)
Value_salt_peper_max=PSNR(original,Max_salt_peper)
Value_salt_peper_min=PSNR(original,Min_salt_peper)
Value_salt_peper_median=PSNR(original,Median_salt_peper)
Value_salt_peper_medianAdapt=PSNR(original,MedianAdapt_salt_peper)

Value_gaussian_media=PSNR(original,media_gaussian)
Value_gaussian_gaussian=PSNR(original,gaussian_gaussian)
Value_gaussian_max=PSNR(original,Max_gaussian)
Value_gaussian_min=PSNR(original,Min_gaussian)
Value_gaussian_median=PSNR(original,Median_gaussian)
Value_gaussian_medianAdapt=PSNR(original,MedianAdapt_gaussian)
```

<!-- grade_id: cell-f472dd5e541ef933, points: 0.4, grade: true, solution: false -->
```python
arreglo=np.arange(0,9).reshape((3,3))
assert np.isclose(PSNR(arreglo,arreglo[::-1]),-13.80,0.1), f'Su función de PSNR no funciona como debería, el PSNR de esta muestra debería ser 34.32 y es {PSNR(arreglo,arreglo[::-1])}'
assert Value_salt_peper_median>Value_salt_peper_media, 'El filtro mediano en sal y pimienta no es el mejor filtro'
assert Value_gaussian_median>Value_gaussian_max, 'El filtro mediano en ruido gaussiano no es el mejor filtro'
assert Value_gaussian_median>Value_gaussian_medianAdapt, 'El filtro mediano en ruido gaussiano no es el mejor filtro'

assert np.isclose(Value_salt_peper_media,18.06,0.1), 'Filtro mal su volumen de sal y pimienta con el filtro medio'
assert np.isclose(Value_salt_peper_gaussian,19.76,0.1), 'Filtro mal su volumen de sal y pimienta con el filtro gaussiano'
assert np.isclose(Value_salt_peper_max,4.12,0.1), 'Filtro mal su volumen de sal y pimienta con el filtro maximo'
assert np.isclose(Value_salt_peper_median,27.94,0.1), 'Filtro mal su volumen de sal y pimienta con el filtro mediano'
assert np.isclose(Value_gaussian_median,26.61,0.1), 'Filtro mal su volumen gaussiano con el filtro mediano'
assert np.isclose(Value_salt_peper_medianAdapt,32.38,0.1), 'Filtro mal su volumen sal y pimienta con el filtro mediano adaptativo'
```

### Parte 4: Segmentación
Ahora, determinaremos cual será nuestro baseline este miniproyecto para segmentar el higado. Lo primero que debemos determinar será nuestra metodología de evaluación.

#### Parte 4.1: Jaccard 3D
Nuestra segmentación no es una imagen binaria, recuerde que estamos tratando con volumenes. En esta ocasión tendremos un volumen de segmentación. El cambio de output y de groundtruth requiere un cambio de métrica de evaluación, ya no podemos usar un Jaccard en 2D como siempre hemos hecho, tampoco podemos hacer un promedio de Jaccards en 2D. Implemente una función de Jaccard 3D que ingrese como parametro el volumen de predicción y el volumen de groundtruth y retorne el indice de Jaccard.

> **Recuerde:** La segmentación de groundtruth tiene 3 diferentes labels, usted quiere una segmentación binaria. Considere el label 1 y 2 como uno mismo.

Cree una función que retorne la precisión, la cobertura y la F-medida dado un volumen de predicción y uno de anotación. Tenga en cuenta los problemas con la división entre 0 (tome epsilon como 1x10^-6)

<!-- grade_id: cell-ac36f0ee7b5d04e6, points: 0, grade: false, solution: true -->
```python
def Jaccard_3D(Im,Gt):
    '''
    :param Im: Volumen de prediccion.
    :param Gt: Segmentacion groundtruth.
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
    return J

def P_C_F_3D(Im,Gt):
    '''
    :param Im: Volumen de prediccion.
    :param Gt: Segmentacion groundtruth.
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
    return Precision,Recall,Fscore
```

<!-- grade_id: cell-0e876058ee00d6bd, points: 0.4, grade: true, solution: false -->
```python
assert np.isclose(Jaccard_3D(original_vol>0.8,seg_vol),0.35,0.01), 'El Jaccard 3D no da los resultados que debe, recuerde que no es lo mismo que un Jaccard 2D promediado'
P,C,F=P_C_F_3D(original_vol>0.8,seg_vol)
assert P<C, 'La precisión debe ser menor que la cobertura'
assert F<C and F>P, 'La F medida debe ser menor que la cobertura y mayor que la precisión'
assert np.isclose(F,0.52,0.1), 'La F medida esta mal calculada'
```

#### Parte 4.2: Elemento conexo
La metodología que utilizaremos para segmentar se basará en las intensidades, sin embargo, muchos tejidos tienen intensidades similares a nuestro organo de interes, lo que resulta en muchos falsos positivos. La idea será umbralizar nuestro volumen, posteriormente nos quedaremos con el componente conexo más grande. Cree una función que reciba como parametro una imagen 2D, la umbralice con un umbral que entre por parametro y retorne la mascara binaria con el elemento conexo mayor.

> - Para determinar el mayor elemento conexo, use skimage.measure.label. Esta función le dará una label a cada elemento conexo, es decir que cambiará su intensidad por un numero de 1 a N siendo N la cantidad de elementos conexos. Puede usar np.sum junto con algun operador logico para determinar el area de cada elemento conexo. La mayor label no necesariamente es el mayor elemento conexo. El fondo es el mayor elemento conexo, no lo tenga en cuenta!

<!-- grade_id: cell-c4ea9649e8963f87, points: 0, grade: false, solution: true -->
```python
from skimage.measure import label

def max_component(Im,umbral,operador):
    'Im (ndarray): Imagen a segmentar'
    'umbral (float): Umbral a utilizar'
    'operador (str): Puede ser "mayor" para > o "menor" para < segun se quiera umbralizar'
    # YOUR CODE HERE
    raise NotImplementedError()
    return max_comp
```

<!-- grade_id: cell-7d4f8ab2f485e83e, points: 0.3, grade: true, solution: false -->
```python
higado_umbral=original_vol[:,:,65]>0.8
higado_seg=max_component(original_vol[:,:,65],0.8,'mayor')
               
fig,ax=plt.subplots(1,2)
ax[1].imshow(higado_umbral)
ax[1].axis('off')
ax[0].set_title('Maximo componente conexo')
ax[0].imshow(higado_seg)
ax[0].axis('off')
ax[1].set_title('Umbralización 0.8')
plt.show()

assert np.sum((1-higado_umbral)*higado_seg)==0, 'Usted segmentó el fondo'
assert np.isclose(np.sum(higado_seg),34863,0.01), 'El elemento conexo segmentado no es el maximo'
```

Ahora, cree una función que entre como parametro un volumen completo y segmente slice por slice (recorrido axial) usando la función max_component realizada anteriormente. 

<!-- grade_id: cell-1bd2c80529d427d0, points: 0, grade: false, solution: true -->
```python
def Vol2D_max_component(Vol,umbral,operador):
    # YOUR CODE HERE
    raise NotImplementedError()
    return segmentacion
```

<!-- grade_id: cell-8cc9ba37b2ac4bc4, points: 0.3, grade: true, solution: false -->
```python
higado_seg_new=Vol2D_max_component(original_vol,0.8,'mayor')
               
fig,ax=plt.subplots(1,3)
ax[2].imshow(higado_seg_new[:,:,65])
ax[2].axis('off')
ax[2].set_title('Slice 65')
ax[0].imshow(higado_seg_new[:,:,40])
ax[0].axis('off')
ax[0].set_title('Slice 40')
ax[1].imshow(higado_seg_new[:,:,55])
ax[1].axis('off')
ax[1].set_title('Slice 55')
plt.show()

assert not np.sum(higado_seg_new[:,:,65]-higado_seg), 'La segmentacion del volumen en el slice 65 debe ser igual a la segmentación independiente del slice'
assert np.isclose(Jaccard_3D(higado_seg_new,seg_vol),0.62,0.05), 'La segmentación de todo el volumen fue erronea'
```

Utilizando esta metodología de segmentación, encuentre el mejor umbral para obtener la mejor segmentación posible. Pongalo en su reporte junto con algunos ejemplos buenos y malos de su implementación. Busque las partes donde el higado se segmenta bien y partes donde no se segmenta el higado.

```python

```

### Parte 5: Visualización de errores
Antes de realizar cualquier pre-procesamiento es importante entender nuestras imagenes y las limitaciones de nuestro modelo para poder diseñar una metodología eficaz que mejore nuestros resultados y no nos resulte perjudicial. Tampoco nos funcionará realizar una metodología de pre-procesamiento a ciegas, siempre es importante comprender nuestro modelo y nuestras imagenes para no solo diseñar un metodo que nos funcione en una base de datos en especifico, sino que sea los mas generalizable posible. 


Cargue el volumen segmentation-1.nii y el volumen volume-1.nii de la base de datos de entrenamiento, luego utilice la función de umbralización creada en la entrega anterior para generar una predicción. Por ultimo, visualize las estructuras de el **volumen original** que usted esta segmentando con su algoritmo. Haga una función que cree 2 volumenes adicionales, uno que contenga unicamente los pixeles falsos positivos y otro que contenga unicamente los pixeles falsos negativos. Para el subplot encuentre en cada volumen los slices (axial, transversal y sagital) con mayor cantidad de errores.

> **Nota:** Recuerde que los FP serán aquellos pixeles que se encuentran en su prediccion pero no en su groundtruth. Los FN serán los pixeles que se encuentran en su groundtruth pero no en su predicción.

<!-- grade_id: cell-0a370a4b035a403c, points: 0, grade: false, solution: true -->
```python
import nibabel as nib
from matplotlib import pyplot as plt
import numpy as np
import os
from skimage.exposure import equalize_hist,match_histograms
from skimage.filters import threshold_otsu as otsu
import cv2
from skimage.measure import label

volume1='' #volumen 1 de entrenamiento
segmen1='' #segmentacion del volumen 1 de entrenamiento, groundtruth
segment_pred='' #volumen de prediccion de su algoritmo
FN_volumes='' #volumen unicamente con pixeles Falsos Negativos
FP_volumes='' #volumen unicamente con pixeles Falsos Positivos

def errores(segment_pred,segmen1):
    'segment_pred (ndarray): volumen de predicción'
    'segmen1 (ndarray): volumen de groundtruth'
    # YOUR CODE HERE
    raise NotImplementedError()
    return FP_volumes,FN_volumes
    
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-79264a30db6c2d76, points: 0.1, grade: true, solution: false -->
```python
aux_segment=Vol2D_max_component(volume1,0.8,'mayor')
fp,fn=errores(aux_segment,segmen1)
assert np.isclose(np.sum(fp),663620, 0.01), f'La cantidad de FP deberian ser 311520 y son {np.sum(fp)}'
assert np.isclose(np.sum(fn),33182, 0.01), f'La cantidad de FN deberian ser 16517 y son {np.sum(fn)}'
```

```python

```

## Informe:
En esta ocasión, no responderemos unicamente las preguntas sino que comenzaremos a familiarizarnos con la estructura de un informe de visión por computador. En este miniproyecto debe hacer un informe con Introducción, Resultados y Conclusiones respondiendo las preguntas en las secciones donde usted crea pertinente responderlas. 

### Exploración de base de datos:
Posteriormente, vamos a utilizar una base de datos que se compone de 3 volumes en train, 1 volumen en valid y 1 volumen en test. Basandonos en esto responda, dentro del informe, las siguientes preguntas:
* ¿Puede afirmar que este problema de segmentación es un problema desbalanceado? ¿Cómo puede afectar esto en nuestra metodología de segmentación?
* Teniendo en cuenta la cantidad de datos ¿Cree usted que un analisis pixel a pixel es eficiente para realizar una segmentación de todos nuestros datos?
* ¿Tiene sentido intentar balancear la base de datos?

### Anotación
* ¿Que particularidad observa entre los histogramas de los volumenes corruptos, el volumen filtrado y el volumen original?
* Teniendo en cuenta los histogramas, que umbrales escogeria en cada caso (volumen original y volumen corrupto) Justifique su respuesta y diga las posibles limitaciones.

### Remoción de ruido:
* ¿Que tipo de ruido es más comun en las imagenes medicas?
* ¿Cuales filtros dañaron más su imagen? ¿Por que sucede esto?
* ¿Cree usted que el filtrado gaussiano o filtro medio afectará su segmentación? Responda teniendo en cuenta nuestra metodología de segmentación y las caracteristicas del filtro.

### Segmentación:
* ¿Porque la metrica de Jaccard 3D no es igual que la de Jaccard 2D promediado? De una explicación matemática.
* Teniendo en cuenta las metricas de Jaccard 3D y la precisión y cobertura responda ¿Porque es necesario establecer una métrica de evaluación acorde a nuestro probelma principal? ¿Son estas 2 métricas compatibles? Conteste esta pregunta teniendo en cuenta el desbalance de datos.
* ¿Que organos o estructuras del cuerpo son problematicas para nuestra metodología de segmentación?

### Visualización de errores
* ¿Que estructuras en la MRI son las mas problematicas para nuestra segmentación?
* ¿Cómo se pueden mejorar los resultados obtenidos mediante una umbralización pura?

```python

```